Overflow metabolism originates from growth optimization and cell heterogeneity

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public Review):

Summary:

Cell metabolism exhibits a well-known behavior in fast-growing cells, which employ seemingly wasteful fermentation to generate energy even in the presence of sufficient environmental oxygen. This phenomenon is known as Overflow Metabolism or the Warburg effect in cancer. It is present in a wide range of organisms, from bacteria and fungi to mammalian cells.

In this work, starting with a metabolic network for Escherichia coli based on sets of carbon sources, and using a corresponding coarse-grained model, the author applies some well-based approximations from the literature and algebraic manipulations. These are used to successfully explain the origins of Overflow Metabolism, both qualitatively and quantitatively, by comparing the results with E. coli experimental data.

By modeling the proteome energy efficiencies for respiration and fermentation, the study shows that these parameters are dependent on the carbon source quality constants K_i (p.115 and 116). It is demonstrated that as the environment becomes richer, the optimal solution for proteome energy efficiency shifts from respiration to fermentation. This shift occurs at a critical parameter value K_A(C).

This counterintuitive result qualitatively explains Overflow Metabolism.

Quantitative agreement is achieved through the analysis of the heterogeneity of the metabolic status within a cell population. By introducing heterogeneity, the critical growth rate is assumed to follow a Gaussian distribution over the cell population, resulting in accordance with experimental data for E. coli. Overflow metabolism is explained by considering optimal protein allocation and cell heterogeneity.

The obtained model is extensively tested through perturbations: 1) Introduction of overexpression of useless proteins; 2) Studying energy dissipation; 3) Analysis of the impact of translation inhibition with different sub-lethal doses of chloramphenicol on Escherichia coli; 4) Alteration of nutrient categories of carbon sources using pyruvate. All model perturbation results are corroborated by E. coli experimental results.

We appreciate the reviewer's highly positive comments and the accurate summary of our manuscript.

Strengths:

In this work, the author employs modeling methods typical of Physics to address a problem in Biology, standing at the interface between these two scientific fields. This interdisciplinary approach proves to be highly fruitful and should be further explored in the literature. The use of Escherichia coli as an example ensures that all hypotheses and approximations in this study are well-founded in the literature. Examples include the approximation for the Michaelis-Menten equation (line 82), Eq. S1, proteome partition in Appendix 1.1 (lines 68-69), and a stable nutrient environment in Appendix 1.1 (lines 83-84). The section "Testing the model through perturbation" heavily relies on bacterial data. The construction of the model and its agreement with experimental data are convincingly presented.

We appreciate the reviewer's highly positive comments. We have incorporated many of the reviewer's insightful suggestions and added citations in the appropriate contexts, which have significantly improved our manuscript.

Weaknesses:

In Section Appendix 6.4, the author explores the generalization of results from bacteria to cancer cells, adapting the metabolic network and coarse-grained model accordingly. It is argued that as a consequence, all subsequent steps become immediately valid. However, I remain unconvinced, considering the numerous approximations used to derive the equations, which the literature demonstrates to be valid primarily for bacteria. A more detailed discussion about this generalization is recommended. Additionally, it is crucial to note that the experimental validation of model perturbations heavily relies on E. coli data.

We appreciate the reviewer's insightful suggestions. We apologize for not clearly illustrating the generalization of results from bacteria to cancer cells in the previous version of our manuscript. Indeed, in our earlier version, there was no experimental validation of model results related to cancer cells.

Following the reviewer’s suggestions, we have now added Fig. 5 and Appendix-fig. 5, fully expanded the previous Appendix 6.4 into Appendix 9 in our current version, and added a new section entitled “Explanation of the Crabtree effect in yeast and the Warburg effect in cancer cells” in our main text to provide a detailed discussion of the generalization from bacteria to yeast and cancer cells. Through the derivations shown in Appendix 9 (Eqs. S180-S189), we arrived at Eq. 6 (or Eq. S190 in Appendix 9) to facilitate the comparison of our model results with experimental data in yeast and cancer cells. This comparison is presented in Fig. 5, where we demonstrate that our model can quantitatively explain the data for the Crabtree effect in yeast and the Warburg effect in cancer cells (related experimental data references: Shen et al., Nature Chemical Biology 20, 1123–1132 (2024); Bartman et al., Nature 614, 349-357 (2023)). These additions have significantly strengthened our manuscript.

Reviewer #2 (Public Review):

Summary

This paper has three parts. The first part applied a coarse-grained model with proteome partition to calculate cell growth under respiration and fermentation modes. The second part considered single-cell variability and performed population average to acquire an ensemble metabolic profile for acetate fermentation. The third part used model and simulation to compare experimental data in literature and obtained substantial consistency.

We thank the reviewer for the accurate summary and positive comments on our manuscript.

Strengths and major contributions

(i) The coarse-grained model considered specific metabolite groups and their interrelations and acquired an analytical solution for this scenario. The "resolution" of this model is in between the Flux Balanced Analysis/whole-cell simulation and proteome partition analysis.

(ii) The author considered single-cell level metabolic heterogeneity and calculated the ensemble average with explicit calculation. The results are consistent with known fermentation and growth phenomena qualitatively and can be quantitatively compared to experimental results.

We appreciate the reviewer’s highly positive comments.

Weaknesses

(i) If I am reading this paper correctly, the author's model predicts binary (or "digital") outcomes of single-cell metabolism, that is, after growth rate optimization, each cell will adopt either "respiration mode" or "fermentation mode" (as illustrated in Figure Appendix - Figure 1 C, D). Due to variability enzyme activity k_i^{cat} and critical growth rate λ_C, each cell under the same nutrient condition could have either respiration or fermentation, but the choice is binary.

The binary choice at the single-cell level is inconsistent with our current understanding of metabolism. If a cell only uses fermentation mode (as shown in Appendix - Figure 1C), it could generate enough energy but not be able to have enough metabolic fluxes to feed into the TCA cycle. That is, under pure fermentation mode, the cell cannot expand the pool of TCA cycle metabolites and hence cannot grow.

This caveat also appears in the model in Appendix (S25) that assumes J_E = r_E*J_{BM} where r_E is a constant. From my understanding, r_E can be different between respiration and fermentation modes (at least for real cells) and hence it is inappropriate to conclude that cells using fermentation, which generates enough energy, can also generate a balanced biomass.

We thank the reviewer for raising this question. Indeed, regarding energy biogenesis between respiration and fermentation, our model predicts binary outcomes at the single-cell level. However, this outcome does not hinder cell growth, as there are three independent possible fates for the carbon source (e.g., glucose) in metabolism: fermentation, respiration for energy biogenesis, and biomass generation. Each fate is associated with a distinct fraction of the proteome, with no overlap between them (see Appendix-figs. 1 and 5). Consequently, in a purely fermentative mode, a cell can still use the proteome dedicated to the biomass generation pathway to produce biomass precursors via the TCA cycle.

The classification of the carbon source’s fates into three independent pathways was initially introduced by Chen and Nielsen (Chen and Nielsen, PNAS 116, 17592-17597 (2019)). We apologize for the oversight in not citing their paper in this context in the previous version of our manuscript (although it was cited elsewhere). We have now included the citation in all appropriate places.

To illustrate this issue more clearly, we explicitly present the proteome allocation results for optimal growth in a fermentation mode below, where the proteome efficiency (i.e., the proteome energy efficiency in our previous version) in fermentation is higher than in respiration (i.e., ). We use the model shown in Fig. 1B as an example, with the relevant equations being Eqs. S26 and S28 in Appendix 2.1. By substituting Eq. S28 into Eq. S26, we arrive at Eq. 3 (or Eq. S29 in Appendix 2.1), which we restate here as Eq. R1:

For a given nutrient condition, i.e., for a specific value of κA at the single-cell level, the values of are determined (see Eqs. S20, S27, S31 and S32), while ϕ and φmax are constants (see Eq. S33 and Appendix 1.1). Therefore, if , then , since all coefficients are positive (i.e., ) and takes non-negative values. Hence, the solution for optimal growth is (see Eqs. S35-S36 in Appendix 2.2):

Here, the result signifies a pure fermentation mode with no respiration flux for energy biogenesis. Then, by combining Eq. R2 with Eqs. S28 and S30 from Appendix 2.1, we obtain the optimal proteome allocation results for this case:

where , while κA and take given values (see Eqs. S20 and S27). In Eq. R3, φ3 corresponds to the fraction of the proteome devoted to carrying the carbon flux from Acetyl-CoA (the entry point of Pool b, see Fig. 1B and Appendix 1.2) to α-Ketoglutarate (the entry point of Pool c), with all of these being enzymes within the TCA cycle. The optimal growth solution is , which demonstrates that in a pure fermentation mode, the optimal growth condition includes the presence of enzymes within the TCA cycle capable of carrying the flux required for biomass generation.

Regarding Eq. S25, JE represents the energy demand for cell proliferation, expressed as the stoichiometric energy flux in ATP. Although the influx of carbon sources (e.g., glucose) varies significantly between fermentation and respiration modes, JBM and JE are the biomass and energy fluxes used to build cells, respectively. In bacteria, whether in fermentation or respiration mode, the proportion of maintenance energy used for protein degradation is roughly negligible (see Locasale and Cantley, BMC Biol 8, 88 (2010)). Consequently, the energy demand represented by _J_E scales approximately linearly with the biomass production rate JBM (related experimental data reference: Ebenhöh et al., Life 14, 247 (2024)), regardless of the energy biogenesis mode. Therefore, _r_E can be regarded as roughly constant for bacteria. However, in eukaryotic cells such as yeast and mammalian cells, the proportion of maintenance energy is much more significant (see Locasale and Cantley, BMC Biol 8, 88 (2010)). Therefore, we have explicitly considered the contribution of maintenance energy in these cases and have extended the previous Appendix 6.4 into Appendix 9 in the current version.

(ii) The minor weakness of this model is that it assumes a priori that each cell chooses its metabolic strategy based on energy efficiency. This is an interesting assumption but there is no known biochemical pathway that directly executes this mechanism. In evolution, growth rate is more frequently considered for metabolic optimization. In Flux Balanced Analysis, one could have multiple objective functions including biomass synthesis, energy generation, entropy production, etc. Therefore, the author would need to justify this assumption and propose a reasonable biochemical mechanism for cells to sense and regulate their energy efficiency.

We thank the reviewer for raising this question and apologize for not explaining this point clearly enough in the previous version of our manuscript. Just as the reviewer mentioned, growth rate should be considered for metabolic optimization under the selection pressure of the evolutionary process. In fact, in our model, the sole optimization objective is exactly the cell growth rate. The determination of whether to use fermentation or respiration based on proteome efficiency (i.e., the proteome energy efficiency in our previous version) is not an a priori assumption in our model; rather, it is a natural consequence of growth rate optimization, as we detail below.

For a given nutrient condition with a determined value of κA , as we have explained in the aforementioned responses, the constraint on the fluxes is summarized in Eq. 3 and is restated as Eq. R1. Mathematically, we can obtain the solution for the optimal growth strategy by combining Eq. R1 (i.e., Eq. 3) with the optimization on cell growth rate λ, and the solution can be obtained as follows: If the proteome efficiency in fermentation is larger than that in respiration, i.e., , then from Eq. R1, we obtain , since the values of εr , εf , Ψ, ϕ and φmax are all fixed for a given κA , with εr , εf , Ψ, ϕ, φmax > 0 . Hence, (since ), and note that . Therefore is the solution for optimal growth, where the growth rate can take the maximum value of . Similarly, for the case where the proteome efficiency in respiration is larger than that in fermentation (i.e ), is the solution for optimal growth. With this analysis, we have demonstrated that the choice between fermentation and respiration based on proteome efficiency is a natural consequence of growth rate optimization.

We have now revised the related content in our manuscript to clarify this point.

My feeling is that the mathematical structure of this model could be correct, but the single-cell interpretation for the ensemble averaging has issues. Each cell could potentially adopt partial respiration and partial fermentation at the same time and have temporal variability in its metabolic mode as well. With the modification of the optimization scheme, the author could have a revised model that avoids the caveat mentioned above.

We thank the reviewer for raising this question. In fact, in the above two responses, we have addressed the issues raised here, clarifying that the binary mode between respiration and fermentation does not hinder cell growth and that the sole optimization objective is the cell growth rate, as the reviewer suggested. Regarding temporal variability, due to factors such as cell cycle stages and the intrinsic noise arising from stochastic processes, temporal variability in the fermentation or respiration mode is indeed likely. However, at any given moment at the single-cell level, a binary choice between fermentation and respiration is what our model predicts for the optimal growth strategy.

Discussion and impact for the field

Proteome partition models and Flux Balanced Analysis are both commonly used mathematical models that emphasize different parts of cellular physiology. This paper has ingredients for both, and I expect after revision it will bridge our understanding of the whole cell.

We appreciate the reviewer’s very positive comments. We have followed many of the good suggestions raised by the reviewer, and our revised manuscript is much improved as a result.

Reviewer #3 (Public Review):

Summary:

In the manuscript "Overflow metabolism originates from growth optimization and cell heterogeneity" the author Xin Wang investigates the hypothesis that the transition into overflow metabolism at large growth rates actually results from an inhomogeneous cell population, in which every individual cell either performs respiration or fermentation.

We thank the reviewer for carefully reading our manuscript and the accurate summary.

Weaknesses:

The paper has several major flaws. First, and most importantly, it repeatedly and wrongly claims that the origins of overflow metabolism are not known. The paper is written as if it is the first to study overflow metabolism and provide a sound explanation for the experimental observations. This is obviously not true and the author actually cites many papers in which explanations of overflow metabolism are suggested (see e.g. Basan et al. 2015, which even has the title "Overflow metabolism in E. coli results from efficient proteome allocation"). The paper should be rewritten in a more modest and scientific style, not attempting to make claims of novelty that are not supported. In fact, all hypotheses in this paper are old. Also the possiblility that cell heterogeneity explains the observed 'smooth' transition into overflow metabolism has been extensively investigated previously (see de Groot et al. 2023, PNAS, "Effective bet-hedging through growth rate dependent stability") and the random drawing of kcat-values is an established technique (Beg et al., 2007, PNAS, "Intracellular crowding defines the mode and sequence of substrate uptake by Escherichia coli and constrains its metabolic activity"). Thus, in terms of novelty, this paper is very limited. It reinvents the wheel and it is written as if decades of literature debating overflow metabolism did not exist.

We thank the reviewer for both the critical and constructive comments. Following the reviewer’s suggestion, we have revised our manuscript to adopt a more modest style. However, we respectfully disagree with the criticism regarding the novelty of our study, as detailed below.

First, while many explanations for overflow metabolism have been proposed, we have cited these in both the previous and current versions of our manuscript. We apologize for not emphasizing the distinctions between these previous explanations and our study in the main text of our earlier version, though we did provide details in Appendix 6.3. In fact, most of these explanations (e.g., Basan et al., Nature 528, 99-104 (2015); Chen and Nielsen, PNAS 116, 17592-17597 (2019); Majewski and Domach, Biotechnol. Bioeng. 35, 732-738 (1990); Niebel et al., Nat. Metab. 1, 125-132 (2019); Shlomi et al., PLoS Comput. Biol. 7, e1002018 (2011); Varma and Palsson, Appl. Environ. Microbiol. 60, 3724-3731 (1994); Vazquez et al., BMC Syst. Biol. 4, 58 (2010); Vazquez and Oltvai, Sci. Rep. 6, 31007 (2016); Zhuang et al., Mol. Syst. Biol. 7, 500 (2011)) heavily rely on the assumption that cells optimize their growth rate for a given rate of carbon influx under each nutrient condition (or certain equivalents) to explain the growth rate dependence of fermentation flux. However, this assumption—that cell growth rate is optimized for a given rate of carbon influx—is questionable, as the given factors in a nutrient condition are the identity and concentration of the carbon source, rather than the carbon influx itself.

Consequently, in our model, we purely optimize cell growth rate without imposing a special constraint on carbon influx. Our assumption that the given factors in a nutrient condition are the identity and concentration of the carbon source aligns with the studies by Molenaar et al. (Molenaar et al., Mol. Syst. Biol. 5, 323 (2009)), where they specified an identical assumption on page 5 of their Supplementary Information (SI); Scott et al. (Scott et al., Science 330, 1099-1102 (2010)), where the growth rate formula was derived for a culturing condition with a given nutrient quality; and Wang et al. (Wang et al., Nat. Comm. 10, 1279 (2019)), our previous study on microbial growth. Among these three studies, only Molenaar et al. addresses overflow metabolism. However, Molenaar et al. did not consider cell heterogeneity, resulting in their model predictions on the growth rate dependence of fermentation flux being a digital response, which is inconsistent with experimental data.

Furthermore, prevalent explanations such as those by Basan et al. (Basan et al., Nature 528, 99-104 (2015)) and Chen and Nielsen (Chen and Nielsen, PNAS 116, 17592-17597 (2019)) suggest that overflow metabolism originates from the proteome efficiency in fermentation always being higher than in respiration. However, Shen et al. (Shen et al., Nature Chemical Biology 20, 1123–1132 (2024)) recently discovered that the proteome efficiency measured at the cell population level in respiration is higher than in fermentation for many yeast and cancer cells, despite the presence of fermentation fluxes through aerobic glycolysis. This finding clearly contradicts the studies by Basan et al. (2015) and Chen and Nielsen (2019).

Nevertheless, our model may resolve this puzzle by incorporating two important features. First, in our model, the proteome efficiency (i.e., the proteome energy efficiency in our previous version) in respiration is larger than that in fermentation when nutrient quality is low (Eqs. S174-S175 in Appendix 9). Second, and crucially, due to the incorporation of cell heterogeneity in our model, there could be a proportion of cells with higher proteome efficiency in fermentation than in respiration, even when the overall proteome efficiency at the cell population level is higher in respiration than in fermentation. As shown in the newly added Fig. 5A-B, our model results can quantitatively illustrate the experimental data from Shen et al., Nature Chemical Biology 20, 1123–1132 (2024).

Finally, regarding the criticism of the novelty of our hypothesis: As specified in our main text, cell heterogeneity has been widely reported experimentally in both microbes (e.g., Ackermann, Nat. Rev. Microbiol. 13, 497-508 (2015); Bagamery et al., Curr. Biol. 30, 4563-4578 (2020); Balaban et al., Science 305, 1622-1625 (2004); Nikolic et al., BMC Microbiol. 13, 1-13 (2013); Solopova et al., PNAS 111, 7427-7432 (2014); Wallden et al., Cell 166, 729-739 (2016)) and tumor cells (e.g., Duraj et al., Cells 10, 202 (2021); Hanahan and Weinberg, Cell 164, 681-694 (2011); Hensley et al., Cell 164, 681-694 (2016)). However, to the best of our knowledge, cell heterogeneity has not yet been incorporated into theoretical models for explaining overflow metabolism or the Warburg effect. The reviewer mentioned the study by de Groot et al. (de Groot et al., PNAS 120, e2211091120 (2023)) as studying overflow metabolism similarly to our work. We have carefully read this paper, including the main text and SI, and found that it is not directly relevant to either overflow metabolism or the Warburg effect. Instead, their model extends the work of Kussell and Leibler (Kussell and Leibler, Science 309, 2075-2078 (2005)), focusing on bet-hedging strategies of microbes in changing environments.

Regarding the criticism that random drawing of kcat-values is an established technique (Beg et al., PNAS 104, 12663-12668 (2007)), we need to stress that the distribution noise on kcat-values considered in our model is fundamentally different from that in Beg et al. In Beg et al., their model involved 876 reactions (see Dataset 1 in Beg et al.), of which only 109 had associated biochemical experimental data. Thus, their distribution of kcat-values pertains to different enzymes within the same cell. In contrast, we have the mean of the kcat-values from experimental data for each relevant enzymes, with the distribution of kcat-values representing the same enzyme in different cells.

Moreover, the manuscript is not clearly written and is hard to understand. Variables are not properly introduced (the M-pools need to be discussed, fluxes (J_E), "energy coefficients" (eta_E), etc. need to be more explicitly explained. What is "flux balance at each intermediate node"? How is the "proteome efficiency" of a pathway defined? The paper continues to speak of energy production. This should be avoided. Energy is conserved (1st law of thermodynamics) and can never be produced. A scientific paper should strive for scientific correctness, including precise choice of words.

We thank the reviewer for the constructive comments. Following these, we have provided more explicit information and revised our manuscript to enhance readability. In our initially submitted version, the phrase "energy production" was borrowed from Nelson et al. (Nelson et al., Lehninger principles of biochemistry, 2008) and Basan et al. (Basan et al., Nature 528, 99-104 (2015)), and we chose to follow this terminology. We appreciate the reviewer’s suggestion and have now revised the wording to use more appropriate expressions.

The statement that the "energy production rate ... is proportional to the growth rate" is, apart from being incorrect - it should be 'ATP consumption rate' or similar (see above), a non-trivial claim. Why should this be the case? Such statements must be supported by references. The observation that the catabolic power indeed appears to increase linearly with growth rate was made, based on chemostat data for E.coli and yeast, in a recent preprint (Ebenhöh et al, 2023, bioRxiv, "Microbial pathway thermodynamics: structural models unveil anabolic and catabolic processes").

We thank the reviewer for the insightful suggestions. Following these, we have revised our manuscript and cited the suggested reference (i.e., Ebenhöh et al., Life 14, 247 (2024)).

All this criticism does not preclude the possibility that cell heterogeneity plays a role in overflow metabolism. However, according to Occam's razor, first the simpler explanations should be explored and refuted before coming up with a more complex solution. Here, it means that the authors first should argue why simpler explanations (e.g. the 'Membrane Real Estate Hypothesis', Szenk et al., 2017, Cell Systems; maximal Gibbs free energy dissipation, Niebel et al., 2019, Nature Metabolism; Saadat et al., 2020, Entropy) are not considered, resp. in what way they are in disagreement with observations, and then provide some evidence of the proposed cell heterogeneity (are there single-cell transcriptomic data supporting the claim?).

We thank the reviewer for raising these questions and providing valuable insights. Regarding the shortcomings of simpler explanations, as explained above, most proposed explanations (including the references mentioned by the reviewer: Szenk et al., Cell Syst. 5, 95-104 (2017); Niebel et al., Nat. Metab. 1, 125-132 (2019); Saadat et al., Entropy 22, 277 (2020)) rely heavily on the assumption that cells optimize their growth rate for a given rate of carbon influx under each nutrient condition (or its equivalents). However, this assumption is questionable, as the given factors in a nutrient condition are the identities and concentrations of the carbon sources, rather than the carbon influx itself.

Specifically, Szenk et al. is a perspective paper, and the original “membrane real estate hypothesis” was proposed by Zhuang et al. (Zhuang et al., Mol. Syst. Biol. 7, 500 (2011)). Zhuang et al. specified in Section 7 of their SI that their model’s explanation of the experimental results shown in Fig. 2C of their manuscript relies on the assumption of restrictions on carbon influx. In Niebel et al. (Niebel et al., Nat. Metab. 1, 125-132 (2019)), the Methods section specifies that the glucose uptake rate was considered a given factor for a growth condition. In Saadat et al. (Saadat et al., Entropy 22, 277 (2020)), Appendix A notes that their model results depend on minimizing carbon influx for a given growth rate, which is equivalent to the assumption mentioned above (see Appendix 6.3 in our manuscript for details).

Regarding the experimental evidence for our proposed cell heterogeneity, Bagamery et al. (Bagamery et al., Curr. Biol. 30, 4563-4578 (2020)) reported non-genetic heterogeneity in two subpopulations of Saccharomyces cerevisiae cells upon the withdrawal of glucose from exponentially growing cells. This strongly indicates the coexistence of fermentative and respiratory modes of heterogeneity in S. cerevisiae cultured in a glucose medium (refer to Fig. 1E in Bagamery et al.). Nikolic et al. (Nikolic et al., BMC Microbiol. 13, 1-13 (2013)) reported a bimodal distribution in the expression of the acs gene (the transporter for acetate) in an E. coli cell population growing on glucose as the sole carbon source within the region of overflow metabolism (see Fig. 5 in Nikolic et al.), indicating the cell heterogeneity we propose. For cancer cells, Duraj et al. (Duraj et al., Cells 10, 202 (2021)) reported a high level of intra-tumor heterogeneity in glioblastoma using optical microscopy images, where 48%~75% of the cells use fermentation and the remainder use respiration (see Fig. 1C in Duraj et al.), which aligns with the cell heterogeneity picture of aerobic glycolysis predicted by our model.

We have now added related content to the discussion section to strengthen our manuscript.

Reviewer #1 (Recommendations For The Authors):

Some minor corrections:

(1) Adjusted the reference: (García-Contreras et al., 2012)

(2) Corrected line 255: Removed the duplicate "the genes"

We thank the reviewer for the suggestions and have implemented each of them to revise our manuscript. The reference in the form of García-Contreras et al., 2012, although somewhat unusual, is actually correct, so we have kept it unchanged.

General comment to the author:

Considering that this work exists at the interface between Physics and Biology, where a significant portion of the audience may not be familiar with the mathematical manipulations performed, it would enhance the paper's readability to provide more explicit indications in the text. For example, in line 91, explicitly define phi_A as phi_R; or in line 115, explain the K_i parameter in the text for better readability.

We thank the reviewer for the suggestion. Following this, we have now provided more explicit information for the definition of mathematical symbols to enhance readability.

Reviewer #2 (Recommendations For The Authors):

The current form of this manuscript is difficult to read for general readers. In addition, the model description in the Appendix can be improved for biophysics readers to keep track of the variables. Here are my suggestions:

a) In the main text, the author should give the definition of "proteome energy efficiency" explicitly both in English and mathematical formula - since this is the central concept of the paper. The biological interpretation of formula (4) should also be stated.

We thank the reviewer for the suggestion. Following this, we have now added definitions and biological interpretations to fix these issues.

b) I feel the basic model of the reaction network in the Appendix could be stated in a more concise way, by emphasizing whether a variable is extensive (exponential growing) or intensive (scale-invariant under exponential growth).

From my understanding, this work assumes balanced exponential growth and hence there is a balanced biomass vector Y* (a constant unit vector with all components sum to 1) for each cell. The steady-state fluxes {J} are extensive and all have growth rate λ. The proteome partition and relative metabolite fractions are ratios of different components of Y* and hence are intensive.

The normalized fluxes {J^(n)} (with respect to biomass) are a function of Y* and are all kept as constant ratios with each other. They are also intensive.

The biomass and energy production are linear combinations of {J} and hence are extensive and follow exponential growth. The biomass and energy efficiency are ratios between flux and proteome biomass, and hence are intensive.

We thank the reviewer for the insightful suggestion. Following this, we have now added the intensive and extensive information for all relevant variables in the newly added Appendix-table 3.

c) In the Appendix, the author should have a table or list of important variables, with their definition, units, and physiological values under respiration and fermentation.

We thank the reviewer for the very useful suggestion. Following this, we have now added Appendix-table 3 (pages 54-57 in the appendices) to illustrate the symbols used throughout our manuscript, as well as the model variables and parameter settings.

d) Regarding the single-cell variability, the author ignored recent experimental measurements on single-cell metabolism. This includes variability on ATP, NAD(P)H in E. coli, which will be useful background for the readers, see below.

https://pubmed.ncbi.nlm.nih.gov/25283467/

https://pubmed.ncbi.nlm.nih.gov/29391569/

We thank the reviewer for the very useful suggestion. We have now cited these relevant studies in our manuscript.

e) The choice between 100% respiration and 100% fermentation is based on the optimization of proteome energy efficiency, while the intermediate strategies are not favored in this model. This is similar to a concept in control theory called the bang-bang principle. This can be added to the Discussion.

We thank the reviewer for this suggestion. We have reviewed the concept and articles on the bang-bang principle. While the bang-bang principle is indeed relevant to binary choices, it is somewhat distant from the topic of metabolic strategies related to optimal growth. The elementary flux mode (see Müller et al., J. Theor. Biol. 347, 182190 (2014); Wortel et al., FEBS J. 281, 1547-1555 (2014)) is more pertinent to this topic, as it may lead to diauxic microbial growth (another binary metabolic strategy) in microbes grown on a mixture of two carbon sources from Group A (see Wang et al., Nat. Comm. 10, 1279 (2019)). Therefore, we have cited and mentioned only the elementary flux mode (Müller et al., J. Theor. Biol. 347, 182-190 (2014); Wortel et al., FEBS J. 281, 1547-1555 (2014)) in the introduction and discussion sections of our manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: A classic problem in metabolism is that fast-proliferating cells use seemingly wasteful fermentation to generate energy in the presence of sufficient oxygen. This counterintuitive phenomenon, known as overflow metabolism, or the Warburg effect in cancer, is universal across various organisms. Despite extensive research, its origin and function remain unclear. Here, we take Escherichia coli as a typical example and show that overflow metabolism can be understood through growth optimization combined with cell heterogeneity. A model of optimal protein allocation, coupled with heterogeneity in enzyme catalytic rates among cells, quantitatively explains why and how cells make the choice between respiration and fermentation under different nutrient conditions. Our model quantitatively illustrates the growth rate dependence of fermentation flux and enzyme allocation under various perturbations, which is fully validated by experimental results. Our work solves the long-standing puzzle of overflow metabolism and can be broadly used to address heterogeneity-related challenges in metabolism.

Revised version: A classic problem in metabolism is that fast-proliferating cells use seemingly wasteful fermentation for energy biogenesis in the presence of sufficient oxygen. This counterintuitive phenomenon, known as overflow metabolism or the Warburg effect, is universal across various organisms. Despite extensive research, its origin and function remain unclear. Here, we show that overflow metabolism can be understood through growth optimization combined with cell heterogeneity. A model of optimal protein allocation, coupled with heterogeneity in enzyme catalytic rates among cells, quantitatively explains why and how cells choose between respiration and fermentation under different nutrient conditions. Our model quantitatively illustrates the growth rate dependence of fermentation flux and enzyme allocation under various perturbations and is fully validated by experimental results in Escherichia coli. Our work provides a quantitative explanation for the Crabtree effect in yeast and the Warburg effect in cancer cells and can be broadly used to address heterogeneity-related challenges in metabolism.

Introduction

Earlier version: A prominent feature of cancer metabolism is that tumor cells excrete large quantities of fermentation products in the presence of sufficient oxygen (Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). This process, discovered by Otto Warburg in the 1920s (Warburg et al., 1924) and known as the Warburg effect, aerobic glycolysis or overflow metabolism, (Basan et al., 2015; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009) is ubiquitous for fast-proliferating cells across a broad spectrum of organisms (Vander Heiden et al., 2009), ranging from bacteria (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; Neidhardt et al., 1990) and fungi (De Deken, 1966) to mammalian cells (Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). For microbes, cells use standard respiration when nutrients are scarce, while they use the counterintuitive aerobic glycolysis when nutrients are adequate, just analogous to normal tissues and cancer cells, respectively (Vander Heiden et al., 2009). Over the past century, especially with extensive studies in the last two decades (Liberti and Locasale, 2016), various rationales for overflow metabolism have been proposed (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Molenaar et al., 2009; Niebel et al., 2019; Pfeiffer et al., 2001; Shlomi et al., 2011; Vander Heiden et al., 2009; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016). In particular, Basan et al. (Basan et al., 2015) provided a systematic characterization of this process, including various types of perturbations in experiments. However, the origin and function of overflow metabolism still remain unclear (DeBerardinis and Chandel, 2020; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009).

Revised version: A prominent feature of cancer metabolism is that tumor cells excrete large quantities of fermentation products in the presence of sufficient oxygen (Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). This process, discovered by Otto Warburg in the 1920s (Warburg et al., 1924) and known as the Warburg effect, aerobic glycolysis, or overflow metabolism (Basan et al., 2015; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009), is ubiquitous among fast-proliferating cells across a broad spectrum of organisms (Vander Heiden et al., 2009), ranging from bacteria (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; Neidhardt et al., 1990) and fungi (De Deken, 1966) to mammalian cells (Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). For microbes, cells use standard respiration when nutrients are scarce, while they use the counterintuitive aerobic glycolysis when nutrients are adequate, just analogous to normal tissues and cancer cells, respectively (Vander Heiden et al., 2009).

Added in the revised version: Over the past century, and particularly through extensive studies in the last two decades (Liberti and Locasale, 2016), various rationales for overflow metabolism have been proposed (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Molenaar et al., 2009; Niebel et al., 2019; Peebo et al., 2015; Pfeiffer et al., 2001; Shlomi et al., 2011; Vander Heiden et al., 2009; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016; Zhuang et al., 2011). Notably, Basan et al. (Basan et al., 2015) provided a systematic characterization of this process, including various types of experimental perturbations. Currently, prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019) hold that overflow metabolism arises from the proteome efficiency in fermentation being consistently higher than that in respiration. However, recent studies have shown that the measured proteome efficiency in respiration is actually higher than in fermentation for many yeast and cancer cells (Shen et al., 2024), even though these cells generate fermentation products through aerobic glycolysis. This finding (Shen et al., 2024) apparently contradicts the prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019). Furthermore, most explanations (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Niebel et al., 2019; Shlomi et al., 2011; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016; Zhuang et al., 2011) heavily rely on the assumption that cells optimize their growth rate for a given rate of carbon influx (i.e., nutrient uptake rate) under each nutrient condition (or its equivalents). However, this assumption is questionable, as the given factors in a nutrient condition are the identities and concentrations of the carbon sources (Molenaar et al., 2009; Scott et al., 2010; Wang et al., 2019), rather than the carbon influx. Therefore, the origin and function of overflow metabolism still remain unclear (DeBerardinis and Chandel, 2020; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009).

Earlier version: Why have microbes and cancer cells evolved to possess the seemingly wasteful aerobic glycolysis strategy? For unicellular organisms, there is evolutionary pressure (Vander Heiden et al., 2009) to optimize cellular resources for rapid growth (Dekel and Alon, 2005; Edwards et al., 2001; Hui et al., 2015; Li et al., 2018; Scott et al., 2010; Towbin et al., 2017; Wang et al., 2019; You et al., 2013). In particular, it has been shown that cells allocate protein resources for optimal growth (Hui et al., 2015; Scott et al., 2010; Wang et al., 2019; You et al., 2013), and the most efficient protein allocation corresponds to elementary flux mode (Müller et al., 2014; Wortel et al., 2014). In this study, we extend these approaches in a heterogeneous framework to address the puzzle of aerobic glycolysis. We use Escherichia coli as a typical example and show that overflow metabolism can be understood from optimal protein allocation combined with the heterogeneity in enzyme catalytic rates. The optimal growth strategy varies between respiration and fermentation depending on the concentration and type of the nutrient, and the combination with cell heterogeneity results in the standard picture (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006) of overflow metabolism. Our model quantitatively illustrates the growth rate dependence of fermentation/respiration flux and enzyme allocation under various types of perturbations, and can be used to explain the Warburg effect in tumor cells.

Revised version: Why have microbes and cancer cells evolved to possess the seemingly wasteful strategy of aerobic glycolysis? For unicellular organisms, there is evolutionary pressure (Vander Heiden et al., 2009) to optimize cellular resources for rapid growth (Dekel and Alon, 2005; Edwards et al., 2001; Hui et al., 2015; Li et al., 2018; Scott et al., 2010; Towbin et al., 2017; Wang et al., 2019; You et al., 2013). In particular, it has been shown that cells allocate protein resources for optimal growth (Hui et al., 2015; Scott et al., 2010; Wang et al., 2019; You et al., 2013), and the most efficient protein allocation corresponds to elementary flux mode (Müller et al., 2014; Wortel et al., 2014). For cancer cells, disrupting the growth control system and evading immune destruction from the host are prominent hallmarks of their survival (Hanahan and Weinberg, 2011), which in certain ways mimic the evolutionary pressure on microbes to optimize cell growth rate. In this study, we apply the optimal growth principle of microbes, which also roughly holds for cancer cells, to a heterogeneous framework to address the puzzle of aerobic glycolysis. We use Escherichia coli as a typical example to show that overflow metabolism can be understood from optimal protein allocation combined with heterogeneity in enzyme catalytic rates. The optimal growth strategy varies between respiration and fermentation depending on the concentration and type of the nutrient, and the combination with cell heterogeneity results in the standard picture (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998) of overflow metabolism. Our model quantitatively illustrates the growth rate dependence of fermentation/respiration flux and enzyme allocation under various types of perturbations in E. coli. Furthermore, it provides a quantitative explanation for the data on the Crabtree effect in yeast and the Warburg effect in cancer cells (Bartman et al., 2023; Shen et al., 2024).

Results

Coarse-grained model

Earlier version: Based on the topology of the metabolic network (Neidhardt et al., 1990; Nelson et al., 2008) (Fig. 1A) , we classify the carbon sources that enter from the upper part of glycolysis into Group A (Wang et al., 2019), and the precursors of biomass components (such as amino acids) into five pools (see Appendix 1.2 for details): a1 (entry point: G6P/F6P), a2 (entry point: GA3P/3PG/PEP), b (entry point: pyruvate/acetyl-CoA), c (entry point: a-ketoglutarate), and d (entry point: oxaloacetate). Pools a1 and a2 are also combined as Pool a due to joint synthesis of precursors. Then, the metabolic network for Group A carbon source utilization (Fig. 1A) is coarse-grained into a model shown in Fig. 1B (see Appendix 2.1 for details) , where node A represents an arbitrary carbon source of Group A. Evidently, Fig. 1B is topologically identical to Fig. 1A. Each coarse-grained arrow in Fig. 1B carries a stoichiometric flux Ji, which delivers carbon flux and may consume or produce energy (e.g., J1, Ja1, see Figs. 1A-B and Appendix-fig. 1A).

Revised version: Based on the topology of the metabolic network (Neidhardt et al., 1990; Nelson et al., 2008) (see Fig. 1A), we classify the carbon sources that enter from the upper part of glycolysis into Group A (Wang et al., 2019) and the precursors of biomass components (such as amino acids) into five pools. Specifically, each pool is designated according to its entry point (see Fig. 1A and Appendix 1.2 for details): a1 (entry point: G6P/F6P), a2 (entry point: GA3P/3PG/PEP), b (entry point: pyruvate/acetyl-CoA), c (entry point: a-ketoglutarate), and d (entry point: oxaloacetate). Pools al and a2 are also combined as Pool a due to the joint synthesis of precursors. Then, the metabolic network for Group A carbon source utilization (see Fig. 1A) can be coarse-grained into a model shown in Fig. 1B (see Appendix 2.1 for details), where node A represents an arbitrary carbon source of Group A. Evidently, Fig. 1B is topologically identical to Fig. 1A. Each coarse-grained arrow in Fig. 1B represents a stoichiometric flux Ji, which delivers carbon flux and may be accompanied by energy consumption or biogenesis (e.g., J1, Ja1, see Figs. 1A-B and Appendix-fig. 1A).

Earlier version: Model and results of overflow metabolism.

Revised version: Model and results of overflow metabolism in E. coli.

Earlier version: (A) The central metabolic network of carbon source utilization. The Group A carbon sources (Wang et al., 2019) are labeled with green squares. (B) Coarse-grained model for Group A carbon source utilization. (C) Model predictions (Eqs. S47 and S160) and experimental results (Basan et al., 2015; Holms, 1996) of overflow metabolism, covering the data for all the Group A carbon sources shown in (A). (D) Growth rate dependence of respiration and fermentation fluxes (Eqs. S47 and S160). (E) The energy efficiencies of respiration and fermentation pathways vary with the growth rate as functions of the substrate quality of a Group A carbon source (Eqs. S31 and S36). See Appendix 8 for model parameter settings and experimental data sources (Basan et al., 2015; Holms, 1996; Hui et al., 2015) of Figs. 1–4.

Revised version: (A) The central metabolic network of carbon source utilization. The Group A carbon sources (Wang et al., 2019) are labeled with green squares. (B) Coarse-grained model for Group A carbon source utilization. (C) Model predictions (see Eqs. S47 and S160) and experimental results (Basan et al., 2015; Holms, 1996) of overflow metabolism, covering the data for all the Group A carbon sources shown in (A). (D) Growth rate dependence of respiration and fermentation fluxes (see Eqs. S47 and S160). (E) The proteome efficiencies for energy biogenesis in the respiration and fermentation pathways vary with growth rate as functions of the substrate quality of a Group A carbon source (see Eqs. S31 and S36). See Appendices 8 and 10 for model parameter settings and experimental data sources (Basan et al., 2015; Holms, 1996; Hui et al., 2015) for Figs. 1–4 of E. coli.

Earlier version: In fact, the stoichiometric flux Ji scales with the cell population. For comparison with experiments, we define the normalized flux Ji(N)≡Ji·m0/Mcarbon , which can be regarded as the flux per unit of biomass (the superscript “(N)” stands for normalized; see Appendix 1.3–1.4 for details). Here, Mcarbon represents the carbon mass of the cell population, and m0 is the weighted average carbon mass of metabolite molecules at the entrance of precursor pools (see Eq. S17). Then, the cell growth rate λ can be represented by λ=∑ia1,a2,b,c,dJi(N) (see Appendix 1.4), and the normalized fluxes of respiration and fermentation are Jr(N)≡J4(N) and Jf(N)≡J6(N) , respectively. In practice, each Ji(N) is characterized by two quantities: the proteomic mass fraction ϕi of the enzyme dedicated to carrying the flux and the substrate quality κi, such that Ji(N)=ϕi·κi . We take the Michaelis-Menten form for the enzyme kinetics, and then κi≡ki·[Si][Si]+Ki . (see Eq. S12 and Appendix 1.4 for details) , where [Si] is the concentration of substrate Si. For each intermediate node and reaction along the pathway (e.g., node M1 in Ja1), the substrate quality κi can be approximated as a constant (see Appendix 1.5): κi≡ki·[Si][Si]+Ki≈ki , where [Si] ≥ Ki generally holds true in bacteria (Bennett et al., 2009; Park et al., 2016). However, κA is a variable that depends on the concentration and nutrient type of a Group A carbon source (see Eq. S27).

Revised version: In fact, the stoichiometric flux Ji scales with the cell population. For comparison with experiments, we define the normalized flux Ji(N)≡Ji⋅m0/Mcarbon , which can be regarded as the flux per unit of biomass (the superscript “(N)” stands for normalized; see Appendix 1.3–1.4 for details). Here, Mcarbon represents the carbon mass of the cell population, and m0 is the weighted average carbon mass of metabolite molecules at the entry of precursor pools (see Eq. S17). Then, the cell growth rate λ can be represented by the total outflow of the normalized fluxes: λ=∑ia1,a2,b,c,dJi(N) (see Appendix 1.4). The normalized fluxes of respiration and fermentation are Jr(N)≡J4(N) and Jf(N)≡J6(N) , respectively (see Fig. 1A-B). In practice, each Ji(N) is characterized by two quantities: the proteomic mass fraction ϕi of the enzyme dedicated to carrying the flux and the substrate quality κi, such that Ji(N)=ϕi⋅κi . We take the Michaelis-Menten form for the enzyme kinetics (Nelson et al., 2008), and then κi≡κi⋅[ Si ][ Si ]+Ki (see Eq. S12 and Appendix 1.4 for details), where [Si] is the concentration of substrate Si, and Ki is the Michaelis constant. For each intermediate node and reaction along the pathway (e.g., node M1 in Ja1), the substrate quality κi can be approximated as a constant (see Appendix 1.5): κi≡κi⋅[ Si ][ Si ]+Ki≈ki , where [Si] ≥ Ki generally holds true in bacteria (Bennett et al., 2009; Park et al., 2016). However, the nutrient quality κA is a variable that depends on the nutrient type and concentration of a Group A carbon source (see Eq. S27).

Earlier version: Generally, there are three distinct destinies of a Group A carbon source in the metabolic network (Appendix-fig. 1C-E) : fermentation, respiration, and biomass generation. Each draws a proteome fraction of ϕf, ϕr and ϕBM. The net effect of the first two destinies is energy production, while the last one generates precursors of biomass accompanied by energy production. By applying the proteomic constraint(Scott et al., 2010) that there is a maximum fraction ϕmax for proteome allocation (ϕmax ≈ 0.48 (Scott et al., 2010)), we have: ϕf+ϕr+ϕBM=ϕmax. (1)

Revised version: Generally, there are three independent fates for a Group A carbon source in the metabolic network (Chen and Nielsen, 2019): fermentation, respiration, and biomass generation (see Appendix-fig. 1C-E). Each draws a distinct proteome fraction of ϕf, ϕr and ϕBM, with no overlap between them (see Appendix 2.1). The net effect of the first two fates is energy biogenesis, while the last one generates precursors for biomass, accompanied by energy biogenesis. By applying the proteomic constraint (Scott et al., 2010) that there is a maximum fraction ϕmax for proteome allocation (ϕmax ≈ 0.48 (Scott et al., 2010)), we have: ϕf+ϕr+ϕBM=ϕmax. (1)

Earlier version: In fact, Eq. 1 is equivalent to ϕR+ϕA+∑j=16ϕj+∑ia1,a2,b,c,dϕi=ϕmax (see Appendix 2.1) , where ϕR stands for the mass fraction of the active ribosome-affiliated proteins. In cell proliferation, ribosomes constitute the majority of enzymes for protein synthesis (Neidhardt et al., 1990; Nelson et al., 2008), while other biomass components such as RNA are optimally produced (Kostinski and Reuveni, 2020) following the growth rate set by protein synthesis. Thus, λ = ϕR · κt, where κt is a parameter determined by the translation rate (Scott et al., 2010) (see Appendix 1.1 for details), which can be approximated as a constant in the growth rate range of concern (Dai et al., 2016).

Revised version: In fact, Eq. 1 is equivalent to ϕR+ϕA+∑j=16ϕj+∑ia1,a2,b,c,dϕi=ϕmax (see Appendix 2.1 for derivation details), where ϕR and ϕA represent the proteomic mass fractions of the active ribosome-affiliated proteins and the cargo proteins responsible for the uptake of the Group A carbon source, respectively. During cell proliferation, ribosomes serve as the factories for protein synthesis and are primarily composed of proteins (Neidhardt et al., 1990; Nelson et al., 2008), while other biomass components, such as RNA, are optimally produced (Kostinski and Reuveni, 2020) in accordance with the growth rate determined by protein synthesis. Thus, the cell growth rate is proportional to ϕR: λ = ϕR ° κt, where κt is a parameter set by the translation rate (Scott et al., 2010) (see Appendix 1.1 for details), which can be approximated as a constant within the growth rate range of interest (Dai et al., 2016).

Added in the revised version: For balanced cell growth in bacteria, the energy demand JE, expressed as the stoichiometric energy flux in ATP, is generally proportional to the biomass production rate (Ebenhoh et al., 2024), since the proportion of maintenance energy is roughly negligible (Locasale and Cantley, 2010) (see Appendix 9 for the cases of yeast and tumor cells). Thus, the normalized flux of energy demand in ATP, denoted as JE(N) , representing the energy demand per unit of biomass, is proportional to the growth rate λ (see Appendix 2.1 for details): JE(N)=ηE⋅λ, (2)

Earlier version: For balanced cell growth, the energy demand is generally proportional to the biomass production rate. Thus, the normalized energy production rate JE(N) is proportional to the growth rate λ (see Appendix 2.1 for details): JE(N)=ηE⋅λ, (2) where ηE is the energy coefficient. By converting all the energy currencies into ATPs, the normalized energy fluxes of respiration and fermentation are Jr(E)=βr(A)·Jr(N)/2 and Jf(E)=βf(A)·Jf(N)/2 , where βr(A) and βf(A) are the stoichiometric coefficients of ATP production per glucose in each pathway (see Appendix-fig. 1C-E and Appendix 2.1 for details). The denominator coefficient (“2”) is derived from the stoichiometry of the coarse-grained reaction M1→M2 (see Fig. 1A-B). Using flux balance analysis at each intermediate node (Mi, i = 1, …, 5) and precursor pool (Pool i, i = a1, a2, b, c, d), combined with the constraints of proteome allocation (Eq. 1) and energy demand (Eq. 2), we obtain the relations between energy fluxes and growth rate for a given nutrient condition with fixed κA (see Appendix 2.1 for details): {Jr(E)+Jf(E)=φ⋅λ,Jr(E)εr+Jf(E)εf=ϕmax−ψ⋅λ, (3) where φ is a constant coefficient mainly determined by ηE (see Eq. S33), and φ·λ represents the normalized energy demand other than the biomass pathway. The coefficients ψ, εr and εf are all functions of κA·. ψ-1 is the proteome efficiency of the biomass pathway (see Eq. S32), with ψ-1 = λ/ϕBM · εr and εf are the proteome energy efficiencies of the respiration and fermentation pathways, with εr≡Jr(E)/ϕr and εf≡Jf(E)/ϕf : {εr=βr(A)1/κA+1/κr(A),εr=βf(A)1/κA+1/κf(A), (4) where both κr(A) and κf(A) are constants, with 1/κr(A) ≡ 1/κ1 + 2/κ2 + 2/κ3 + 2/κ4 and 1/κr(A) ≡ 1/κ1 + 2/κ2 + 2/κ6 (see Appendix 2.1 for details).

Revised version: where ηE is an energy coefficient (see Eqs. S25–S26 for details). By converting all energy currencies (such as NADH, FADH2, etc.) into ATP, the normalized energy fluxes for respiration and fermentation are given by Jr(E)=βr(A)⋅Jr(N)/2 and Jf(E)=βf(A)⋅Jf(N)/2 , where βr(A) and βf(A) are the stoichiometric coefficients of ATP production per glucose in each pathway (see Appendixfig. 1C-E and Appendix 2.1 for details). The denominator coefficient of “2” is derived from the stoichiometry of the coarse-grained reaction M1→2M2 (see Fig. 1A-B). Applying the criteria of flux balance (i.e., mass conservation; see Appendix 1.3) at each intermediate node (Mi, i = 1, …, 5) and precursor pool (Pool i, i = a1, a2, b, c, d), along with the constraints of proteome allocation (see Eq. 1) and energy demand (see Eq. 2), we obtain the relations between normalized energy fluxes and growth rate for a given nutrient condition with a fixed κA (see Appendix 2.1 for details): {Jr(E)+Jf(E)=φ⋅λ,Jr(E)εr+Jf(E)εf=ϕmax−ψ⋅λ, (3)

Added in the revised version: where φ is a constant coefficient primarily determined by ηE (see Eq. S33), and φ·λ represents the normalized flux of energy demand, excluding the energy biogenesis from the biomass synthesis pathway. The coefficients ψ, ψr and ψf are functions of κA · ψ-1 denotes the proteome efficiency for biomass generation in the biomass synthesis pathway (see Eq. S32), defined as ψ-1 = λ/ϕBM (see Appendix 2.1). ψr and ψf represent the proteome efficiencies for energy biogenesis in the respiration and fermentation pathways, respectively, defined as the normalized energy fluxes expressed in ATP generated per proteomic mass fraction, with εr=Jr(E)/ϕr and εf=Jf(E)/ϕf . Hence, {εr=βr(A)1/κA+1/κr(A),εr=βf(A)1/κA+1/κf(A), (4)

Added in the revised version: where both κr(A) and κf(A) are composite parameters that can be approximated as constants, with 1/κr(A) = 1/κ1 + 2/κ2 + 2/κ3 + 2/κ4 and 1/κf(A) = 1/κ1 + 2/κ2 + 2/κ6 (see Appendices 1.5 and 2.1 for details).

Origin of overflow metabolism

Earlier version: The standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006) is exemplified by the experimental data (Basan et al., 2015) shown in Fig. 1C: the fermentation flux exhibits a threshold-analog dependence on the growth rate λ. It is well known that respiration is far more efficient than fermentation in terms of energy production per unit carbon (i.e., βr(A) > βf(A) ) (Nelson et al., 2008; Vander Heiden et al., 2009). Then, why do cells bother to use the seemingly wasteful fermentation pathway? We proceed to resolve these issues using the optimal protein allocation hypothesis (Scott et al., 2010; You et al., 2013).

Revised version: The standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998) is exemplified by the experimental data (Basan et al., 2015) presented in Fig. 1C, where the fermentation flux exhibits a threshold-analog dependence on the growth rate λ. It is well established that respiration is significantly more efficient than fermentation in terms of energy biogenesis per unit of carbon (i.e., βr(A)>βf(A) ) (Nelson et al., 2008; Vander Heiden et al., 2009). Then, why do cells bother to use the seemingly wasteful fermentation pathway? We proceed to address this issue by applying optimal protein allocation (Scott et al., 2010; Wang et al., 2019) within the framework of optimal growth.

Earlier version: For cell proliferation in a given nutrient with fixed κA , the values of εr, εf and ψ are determined (Eqs. 4 and S32). However, the growth rate λ can be influenced by protein allocation between ϕr and ϕf with the governing equation Eq. 3. If εr < εr, then λ=ϕmax−Jf(E)(1/εf−1/εr)ψ+φ/εr≤ϕmaxψ+φ/εr . The optimal solution is εf = Jf(E) = 0, which means that the cell only uses respiration. Conversely, if εf > εr, ϕr = Jr(E) = 0 is optimal, and the cell only uses fermentation. In either case, the optimal choice between respiration and fermentation is determined by weighing the proteome energy efficiencies.

Revised version: For cell proliferation in a given nutrient condition (i.e., with a fixed κA), the values of εr, εf and ψ are determined (Eqs. 4 and S32). However, the growth rate λ can be influenced by protein allocation between respiration and fermentation, specifically λr and λf, according to the governing equation (Eq. 3). If εr > εf, that is, if the proteome efficiency in respiration is higher than that in fermentation, then λ=ϕmax−Jr(E)(1/εf−1/εr)Ψ+φ/εr≤ϕmaxΨ+φ/εr . The optimal growth strategy is ϕf=Jf(E)=0 , meaning that the cell exclusively uses respiration. Conversely, if εf> εr, then ϕr=Jr(E)=0 is optimal, and the cell solely uses fermentation. In either case, the choice between respiration and fermentation for optimal growth is determined by weighing the proteome efficiencies.

Earlier version: In practice, both εr and εf are functions of κA (Eq. 4), and therefore the optimal choice may vary depending on the nutrient conditions. In nutrient-poor conditions where κA ≪ κr(A) and κA ≪ κf(A) , the proteome energy efficiencies are reduced to εr ≈ βr(A) · κA and εf ≈ βf(A) · κA (see Eq. 4), where εr (κA)> εf (κA). In rich media, however, using the parameters of κi derived from in vivo/in vitro experimental data of Escherichia coli (Appendix-table 1, see also Appendix 6.1-6.2 and Appendix-table 2), we confirm that εr (κglucose(ST)) < εf (κglucose(ST)) with Eq. 4 (see also Eqs. S39–S40), where κglucose(ST) represents the substrate quality of glucose at saturated concentration (abbreviated as “ST” in the superscript). Indeed, recent experimental studies have validated that εr (κlactose(ST)) < εf (κlactose(ST)) for lactose (Basan et al., 2015). In Fig. 1E, we present the growth rate dependence of proteome energy efficiencies εr and εf in a three-dimensional (3D) form using the data from Appendix-table 1, where εr, εf and the growth rate λ all vary as functions of κA. Therefore, the ratio εf/εr (defined as Δ(κA)) is a monotonically increasing function of κA, and there exists a critical value of κA (denoted as κA(C) , see Appendix 2.2 for details) satisfying Δ(κA(C))=1 . Below κA(C) , the cell grows slowly and the optimal choice is respiration (with εf < εr, λ = ϕmax ·(ψ + φ/εr)-1), while above κA(C) , the cell grows faster and the optimal growth strategy is fermentation (with εf > εr, λ = ϕmax ·(ψ + φ/εf)-1). The above analysis qualitatively explains the phenomenon of aerobic glycolysis.

Revised version: In practice, both proteome efficiencies εr and εf are functions of nutrient quality κA, which can be influenced by the nutrient type and concentration of the carbon source (see Eqs. 4 and S27). Therefore, the choice for optimal growth may vary depending on the nutrient conditions. In nutrient-poor conditions where κA≪κr(A) and κA≪κf(A) , the proteome efficiencies can be approximated by εr≈βr(A)⋅κA and εf≈βf(A)⋅κA (see Eq. 4), and hence εr (κA)>εf(κA) (since βr(A)>βf(A) ), meaning that the proteome efficiency of respiration is higher than that of fermentation under nutrient-poor conditions. In rich media, however, using the parameters of κi derived from in vivo/in vitro experimental data of E. coli (see Appendix-tables 1–2 and Appendix 6.1–6.2), we obtain εr(κglucose(ST))<εf(κglucose(ST)) with Eq. 4 (see also Eqs. S39–S40), where (κglucose(ST)) represents the substrate quality of glucose at saturated concentration (abbreviated as “ST” in the superscript).

Added in the revised version: This means that the proteome efficiency in fermentation is higher than that in respiration for bacteria in rich media. Indeed, recent studies have validated that the measured proteome efficiency in fermentation is higher than in respiration for E. coli in lactose at saturated concentration (Basan et al., 2015), i.e., εr(κlactose(ST))<εf(κlactose(ST)) . In Fig. 1E, we present the growth rate dependence of proteome efficiencies εr and εf in a three-dimensional (3D) form using the collected data shown in Appendix-table 1, where εr, εf and the growth rate λ all vary as functions of nutrient quality κA. Furthermore, the ratio Δ (defined as Δ(κA) ≡ εf (κA)/εr (κA)) is a monotonically increasing function of κA, and there exists a critical value of κA (denoted as κA(C) ; see Appendix 2.2 for details) satisfying Δ(κA(C))=1 . Below κA(C) , the cell grows slowly, and the choice for optimal growth is respiration (with εf < εr, λ = ϕmax · (ψ + φ/εr)-1), while above κA(C) , the cell grows faster, and the optimal growth strategy is fermentation (with εf < εr, λ = ϕmax · (ψ + φ/εf)-1)). The above analysis qualitatively explains the phenomenon of aerobic glycolysis.

Earlier version: For a quantitative understanding of overflow metabolism, let us first consider the homogeneous case, where all cells share identical biochemical parameters. For optimal protein allocation, the relation between fermentation flux and growth rate is Jf(E) = φ·φ·θ(φ − φC), where “θ” represents the Heaviside step function, and λC is the critical growth rate for κA(C) (i.e., λC ≡ λ (κA(C)) ). Similarly, we can determine the growth rate dependence of respiration flux, which is Jr(E) = φ·λ·[1 − θ(λ − λC)]. These digital response outcomes are consistent with the numerical simulation findings of Molenaar et al. (Molenaar et al., 2009). However, they are clearly incompatible with the threshold-analog response in the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006).

Revised version: For a quantitative understanding of overflow metabolism, let us first consider the homogeneous case, where all cells share identical biochemical parameters. For optimal protein allocation, the relation between fermentation flux and growth rate is given by Jf(E)=φ⋅λ⋅θ(λ−λC) , where “Q” represents the Heaviside step function, and λC is the critical growth rate for a nutrient condition with κA(C) (i.e., λC ≡ λ( κA(C) )). Similarly, the growth rate dependence of respiration flux is Jf(E)=φ⋅λ⋅[ 1−θ(λ−λC) ] . These digital response outcomes are consistent with the numerical simulation findings of Molenaar et al. (Molenaar et al., 2009). However, they are clearly incompatible with the threshold-analog response observed in the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998).

Earlier version: To address this issue, we take into account cell heterogeneity, which is ubiquitous in both microbes (Ackermann, 2015; Bagamery et al., 2020; Balaban et al., 2004; Nikolic et al., 2013; Solopova et al., 2014; Wallden et al., 2016) and tumor cells (Duraj et al., 2021; Hanahan and Weinberg, 2011; Hensley et al., 2016). For the Warburg effect or overflow metabolism of our concern, experimental studies have reported significant metabolic heterogeneity in the choice between respiration and fermentation within a cell population (Bagamery et al., 2020; Duraj et al., 2021; Hensley et al., 2016; Nikolic et al., 2013). Motivated by the fact that the turnover number (kcat value) of a catalytic enzyme varies considerably between in vitro and in vivo measurements (Davidi et al., 2016; García - Contreras et al., 2012), we note that the concentrations of potassium and phosphate, which vary from cell to cell, have a significant impact on the kcat values of the metabolic enzymes (García - Contreras et al., 2012). Therefore, in a cell population, there is a distribution of values for kcat, which is commonly referred to as extrinsic noise (Elowitz et al., 2002). For simplicity, we assume that each kcat value follows a Gaussian distribution. This gives the distributions of single-cell growth rate in various types of carbon sources (see Eqs. S155–S157, S163–S165), which are fully verified by recent experiments using isogenic Escherichia coli with single-cell resolution (Wallden et al., 2016) (Appendix-fig. 2B). Then, the critical growth rate λC should follow a Gaussian distribution N(μλC,σλC2) in the cell population (see Appendix 7 for details), where μλC is approximated by the deterministic result of λC (Eq. S43). We assume that the coefficient of variation (CV) is σλC/μλC = 12%, or equivalently, the CV for the catalytic rate of each metabolic enzyme is 25%. Thus, we have the growth rate dependence of fermentation and respiration fluxes (see Appendix 2.3 for details): {Jf(N)(λ)=φ⋅λβf(A)⋅[erf(λ−μλC2σλC)+1],Jr(N)(λ)=φ⋅λβr(A)⋅[1−erf(λ−μλC2σλC)], (5) where “erf” represents the error function. The fermentation flux exhibits a threshold-analog relation with the growth rate (the red curves in Figs. 1C-D, 2B-C and 3B, D, F), while the respiration flux (the blue curve in Fig. 1D) decreases with an increase in fermentation flux. In Fig. 1C-D, we see that the model results (see Eq. 5 and Appendix 8; parameters are set by the experimental data shown in Appendix-table S1) agree quantitatively with the experimental data of Escherichia coli (Basan et al., 2015; Holms, 1996). The fermentation flux was determined by the acetate secretion rate Jactate(M)=2Jf(N) , and the respiration flux was deduced from the carbon dioxide flux JCO2,r(M)=6Jr(N) (the superscript “(M)” represents the measurable flux in the unit of mM/OD600/h; see Appendix 8.1 for details). Therefore, with the integration of cell heterogeneity, our model of optimal protein allocation quantitatively explains overflow metabolism.

Revised version: To address this issue, we take into account cell heterogeneity, which is ubiquitous in both microbes (Ackermann, 2015; Bagamery et al., 2020; Balaban et al., 2004; Nikolic et al., 2013; Solopova et al., 2014; Wallden et al., 2016; Yaginuma et al., 2014; Zhang et al., 2018) and tumor cells (Duraj et al., 2021; Hanahan and Weinberg, 2011; Hensley et al., 2016). In the context of the Warburg effect or overflow metabolism, experimental studies have reported significant metabolic heterogeneity in the choice between respiration and fermentation within a cell population (Bagamery et al., 2020; Duraj et al., 2021; Hensley et al., 2016; Nikolic et al., 2013). Motivated by the observation that the turnover number (kcat value) of a catalytic enzyme varies considerably between in vitro and in vivo measurements (Davidi et al., 2016; Garcia-Contreras et al., 2012), we note that the concentrations of potassium and phosphate, which vary from cell to cell, have a significant impact on the kcat values of metabolic enzymes (Garcia-Contreras et al., 2012). Therefore, within a cell population, there is a distribution of kcat values for a catalytic enzyme, commonly referred to as extrinsic noise (Elowitz et al., 2002). For simplicity, we assume that the kcat values for each enzyme follow a Gaussian distribution. This leads to varied distributions of single-cell growth rate in different carbon sources (see Eqs. S155–S157, S163–S165), which has been fully verified by recent experiments using isogenic E. coli at single-cell resolution (Wallden et al., 2016) (Appendix-fig. 2B). Consequently, the critical growth rate λC should follow a Gaussian distribution Ν(μλC,σλC2) within a cell population (see Appendix 7 for details), where μλC is approximated by the deterministic result of λC (Eq. S43). We assume that the coefficient of variation (CV) is σλC/μλC = 12%, or equivalently, the CV for the catalytic rate of each metabolic enzyme is 25%. Thus, we obtain the growth rate dependence of fermentation and respiration fluxes (see Appendix 2.3 for details): {Jf(N)(λ)=φ⋅λβf(A)⋅[erf(λ−μλC2σλC)+1],Jr(N)(λ)=φ⋅λβr(A)⋅[1−erf(λ−μλC2σλC)], (5)

Added in the revised version: where “erf” represents the error function. The fermentation flux exhibits a threshold-analog relation with the growth rate (the red curves in Figs. 1C-D, 2B-C, and 3B, D, F), while the respiration flux (the blue curve in Fig. 1D) decreases with an increase in fermentation flux. In Fig. 1C-D, we observe that the model results (see Eq. 5 and Appendix 8; parameters are set based on the experimental data shown in Appendix-table S1) quantitatively agree with the experimental data from E. coli (Basan et al., 2015; Holms, 1996). The fermentation flux was determined by the acetate secretion rate Jactate(M)=2Jf(N) , and the respiration flux was deduced from the carbon dioxide flux JCO2,r(M)=6Jr(N) (the superscript “(M)” represents the measurable flux in the unit of mM/OD600/h; see Appendix 8.1 for details). Therefore, by incorporating cell heterogeneity, our model of optimal protein allocation quantitatively explains overflow metabolism.

Earlier version: Influence of protein overexpression on overflow metabolism.

Revised version: Influence of protein overexpression on overflow metabolism in E. coli.

Earlier version: (A) A 3D plot of the relations among fermentation flux, growth rate, and the expression level of useless proteins. In this plot, both the acetate excretion rate and growth rate vary as bivariate functions of the substrate quality of a Group A carbon source (denoted as κA) and the useless protein expression encoded by LacZ (denoted as ϕZ perturbation, see Eqs. S57 and S160). (B) Growth rate dependence of the acetate excretion rate upon ϕZ perturbation for each fixed nutrient condition (Eq. S58 and S160). (C) Growth rate dependence of the acetate excretion rate as κA varies (Eqs. S57 and S160), with each fixed expression level of LacZ.

Revised version: (A) A 3D plot of the relations among fermentation flux, growth rate, and the expression level of useless proteins. In this plot, both the acetate excretion rate and growth rate vary as bivariate functions of the nutrient quality of a Group A carbon source (denoted as κA) and the useless protein expression encoded by lacZ gene (denoted as ϕZ perturbation; see Eqs. S57 and S160). (B) Growth rate dependence of the acetate excretion rate upon ϕZ perturbation for each fixed nutrient condition (see Eqs. S58 and S160). (C) Growth rate dependence of the acetate excretion rate as κA varies (see Eqs. S57 and S160), with each fixed expression level of LacZ.

Earlier version: Influence of energy dissipation, translation inhibition, and carbon source category alteration on overflow metabolism.

Revised version: Influence of energy dissipation, translation inhibition, and carbon source category alteration on overflow metabolism in E. coli.

Earlier version: (A) A 3D plot of the relations among fermentation flux, growth rate, and the energy dissipation coefficient (Eqs. S70 and S160). (B) Growth rate dependence of the acetate excretion rate as κA varies, with each fixed energy dissipation coefficient determined by/fitted from experimental data. (C) A 3D plot of the relations among fermentation flux, growth rate, and the translation efficiency (Eqs. 85 and S160). Here, the translation efficiency is adjusted by the dose of chloramphenicol (Cm). (D) Growth rate dependence of the acetate excretion rate as κA varies, with each fixed dose of Cm. (E) Coarse-grained model for pyruvate utilization. (F) The growth rate dependence of fermentation flux in pyruvate (Eqs. 105 and S160) significantly differs from that of the Group A carbon sources (Eqs. 47 and S160).

Revised version: (A) A 3D plot of the relations among fermentation flux, growth rate, and the energy dissipation coefficient (see Eqs. S70 and S160). (B) Growth rate dependence of the acetate excretion rate as the nutrient quality κA varies, with each fixed energy dissipation coefficient determined by or fitted from experimental data. (C) A 3D plot of the relations among fermentation flux, growth rate, and the translation efficiency (see Eqs. 85 and S160). Here, the translation efficiency is adjusted by the dose of chloramphenicol (Cm). (D) Growth rate dependence of the acetate excretion rate as κA varies, with each fixed dose of Cm. (E) Coarse-grained model for pyruvate utilization. (F) The growth rate dependence of fermentation flux in pyruvate (see Eqs. 105 and S160) significantly differs from that of the Group A carbon sources (see Eqs. 47 and S160).

Testing the model through perturbations

Earlier version: To further test our model, we systematically investigate model predictions under various types of perturbations and compare them with the experimental data from existing studies (Basan et al., 2015; Holms, 1996) (see Appendices 3 and 4.1 for details).

Revised version: To further test our model, we systematically investigate its predictions under various types of perturbations and compare them with experimental data from existing studies (Basan et al., 2015; Holms, 1996) (see Appendices 3 and 4.1 for details).

Earlier version: First, we consider the proteomic perturbation by overexpression of useless proteins encoded by the Lacz gene (i.e., ϕZ perturbation) in Escherichia coli. The net effect of the ϕZ perturbation is that the maximum fraction of proteome available for resource allocation changes from ϕmax to ϕmax − ϕZ5, where ϕZ is the mass fraction of useless proteins. In a cell population, the critical growth rate λC(ϕZ) still follows a Gaussian distribution N(μλC (ϕZ), σλC (ϕZ)2), where the CV of λC (ϕZ) remains unchanged. Consequently, the growth rate dependence of fermentation flux changes into Jf(N)=φ·λβf(A)·[erf(λ−μλC(ϕZ)2σλC(ϕZ))+1] (see Appendix 3 for model perturbation results of respiration flux), where both the growth rate λ(κA, ϕZ) and the normalized fermentation flux Jf(N) (κA, ϕZ) are bivariate functions of κA and ϕZ (see Eqs. S49 and S56–S57). For each degree of LacZ expression (with fixed ϕZ), similar to the wild-type strains, the fermentation flux exhibits a threshold-analog response to growth rate as κA varies (Fig. 2C), which agrees quantitatively with experimental results (Basan et al., 2015). The shifts in the critical growth rate λC(ϕZ) are fully captured by μλC (ϕZ) = μλC(0)(1 − ϕZ/ϕmax). In contrast, for nutrient conditions with each fixed κA, since the growth rate changes with ϕZ just like λC(ϕZ): λ(κA, ϕZ) = λ(κA, 0)(1 − ϕZ/ϕmax), the fermentation flux is then proportional to the growth rate for the varying levels of Lacz expression: Jf(N)=φβf(A)·[erf(λ(κA,0)−μλC(0)2σλC(0))+1]·λ , where the slope is a monotonically increasing function of the substrate quality κA. These scaling relations are well validated by the experimental data (Basan et al., 2015) shown in Fig. 2B. Finally, in the case where both κA and ϕZ are free to vary, the growth rate dependence of fermentation flux presents a threshold-analog response surface in a 3D plot where ϕZ appears explicitly as the y-axis (Fig. 2A). Experimental data points (Basan et al., 2015) lie right on this surface (Fig. 2A), which are highly consistent with the model predictions.

Revised version: First, we consider the proteomic perturbation caused by overexpression of useless proteins encoded by the lacZ gene (i.e., ϕZ perturbation) in E. coli. The net effect of the ϕZ perturbation is that the maximum fraction of the proteome available for resource allocation changes from ϕmax to ϕmax − ϕZ (Basan et al., 2015), where ϕZ is the proteomic mass fraction of useless proteins. In a cell population, the critical growth rate λC (ϕZ) still follows a Gaussian distribution N (μλC (ϕZ), σλC (ϕZ)2), where the CV of λC (ϕZ) remains unchanged. Consequently, the growth rate dependence of fermentation flux changes to Jf(N)=φ⋅λβf(A)⋅[ erf(λ−μλC(ϕZ)2σλC(ϕZ))+1 ] (see Appendix 3 for model perturbation results regarding respiration flux), where both the growth rate λ(κA, ϕZ) and the normalized fermentation flux Jf(N)(κA,ϕZ) are bivariate functions of κA and ϕZ (see Eqs. S49 and S56–S57). For each degree of LacZ expression (with fixed ϕZ), similar to wild-type strains, the fermentation flux exhibits a threshold-analog response to growth rate as κA varies (see Fig. 2C), which agrees quantitatively with experimental results (Basan et al., 2015). The shifts in the critical growth rate λC(ϕZ) are fully captured by μλC (ϕZ) = μλC (0) (1 − ϕZ/ϕmax). In contrast, for nutrient conditions with each fixed κA, since the growth rate changes with ϕZ just like λC (ϕZ): λ(κA, ϕZ) = λ (κA, 0)(1 − ϕZ/ϕmax), the fermentation flux is then proportional to the growth rate for the varying levels of LacZ expression: Jf(N)=φβf(A)⋅[ erf(λ(κA)−μλC(0)2σλC(0))+1 ]⋅λ , where the slope is a monotonically increasing function of the substrate quality κA. These scaling relations are well validated by the experimental data (Basan et al., 2015) shown in Fig. 2B. Finally, in the case where both κA and ϕZ are free to vary, the growth rate dependence of fermentation flux presents a threshold-analog response surface in a 3D plot, where ϕZ appears explicitly as the y-axis (see Fig. 2A). Experimental data points (Basan et al., 2015) lie right on this surface, which is highly consistent with the model predictions.

Earlier version: Next, we study the influence of energy dissipation, which introduces an energy dissipation coefficient “w” in Eq. 2: JE(N) = ηE · λ + w. Similarly, the critical growth rate λC (w) follows a Gaussian distribution N(μλC(w), σλC(w)2) in a cell population. The relation between the growth rate and fermentation flux can be characterized by (see Appendix 3.2 for details): Jf(N)=φ·λ+wβf(A)·[erf(λ−μλC(w)2σλC(w)2)+1] . In Fig. 3A-B, we present a comparison between the model results and experimental data (Basan et al., 2015) in 3D and 2D plots, which demonstrate good agreement. A notable characteristic of energy dissipation, as distinguished from ϕZ perturbation, is that the fermentation flux increases despite a decrease in the growth rate when κA is fixed.

Revised version: Next, we study the influence of energy dissipation, which introduces an energy dissipation coefficient w to Eq. 2: JE(N)=ηE⋅λ+w . Similarly, the critical growth rate in this case, λC(w), follows a Gaussian distribution N(μλC(w), σλC(w)2) in a cell population. The relation between the growth rate and fermentation flux can be characterized by (see Appendix 3.2 for details): Jf(N)=φ⋅λ+wβf(A)⋅[ erf(λ−μλC(w)2σλC(w)2)+1 ] . In Fig. 3A-B, we present a comparison between the model results and experimental data (Basan et al., 2015) in 3D and 2D plots, which demonstrate good agreement. A notable characteristic of energy dissipation, as distinguished from ϕZ perturbation, is that the fermentation flux increases despite a decrease in the growth rate when κA is fixed.

Earlier version: We proceed to analyze the impact of translation inhibition with different sub-lethal doses of chloramphenicol on Escherichia coli. This type of perturbation introduces an inhibition coefficient “l” in the translation rate, thus turning κt into κt/(l+1). Still, the critical growth rate λC(l) follows a Gaussian distribution N(μλC(l), σλC(l)2), and then, the growth rate dependence of fermentation flux is: Jf(N)=φ·λβf(A)·[erf(λ−μλC(ι)2σλC(ι))+1] (see Appendix 3.3 for details). In Appendix-fig. 2D-E, we see that the model predictions are generally consistent with the experimental data (Basan et al., 2015). However, there is a noticeable systematic discrepancy when the translation rate is small. Therefore, we turn to maintenance energy, which is tiny and generally negligible for bacteria over the growth rate range of concern (Basan et al., 2015; Neidhardt, 1996). Encouragingly, by assigning a very small value to the maintenance energy coefficient “w0” (w0 = 2.5 (h-1)), the model results for the growth rate-fermentation flux relation Jf(N)=φ·λ+w0βf(A)·[erf(λ−μλC(ι)2σλC(ι))+1] agree quantitatively with experiments (Basan et al., 2015) (see Fig. 3C-D and Appendix 3.3 for details).

Revised version: We proceed to analyze the impact of translation inhibition with different sub-lethal doses of chloramphenicol on E. coli. This type of perturbation introduces an inhibition coefficient ι to the translation rate, thus turning κt into κt/(ι + 1). Still, the critical growth rate λC(ι) follows a Gaussian distribution N (μλC(ι), σλC(ι)2), and then, the growth rate dependence of fermentation flux is is given by: Jf(N)=φ⋅λβf(A)⋅[ erf(λ−μλC(ι)2σλC(ι)2)+1 ] (see Appendix 3.3 for details). In Appendix-fig. 2D-E, we observe that the model predictions are generally consistent with the experimental data (Basan et al., 2015). However, a noticeable systematic discrepancy arises when the translation rate is low. Therefore, we consider maintenance energy, which is typically tiny and generally negligible for bacteria over the growth rate range of interest (Basan et al., 2015; Locasale and Cantley, 2010; Neidhardt, 1996). Encouragingly, by assigning a very small value to the maintenance energy coefficient w0 (w0 = 2.5 (h-1)), the model results for the growth rate-fermentation flux relation Jf(N)=φ⋅λ+w0βf(A)⋅[ erf(λ−μλC(ι)2σλC(ι))+1 ] quantitatively agree with experiments (Basan et al., 2015) (see Fig. 3C-D and Appendix 3.3 for details).

Earlier version: Finally, we consider the alteration of nutrient categories by switching to a non-Group A carbon source: pyruvate, which enters the metabolic network from the endpoint of glycolysis (Neidhardt et al., 1990; Nelson et al., 2008). The coarse-grained model for pyruvate utilization is shown in Fig. 3E (see also Fig. 1A), which shares identical precursor pools as those for Group A carbon sources (Fig. 1B), yet with several differences in the coarse-grained reactions. The growth rate dependencies of both the proteome energy efficiencies (Appendix-fig. 2H) and energy fluxes are qualitatively similar to those of the Group A carbon source utilization, while there are quantitative differences in the coarse-grained parameters (see Appendices 4.1 and 8 for details). Most notably, the critical growth rate λC(py) and the ATP production per glucose in the fermentation pathway βf(py) are noticeably smaller than λC and βf(A) (for Group A sources), respectively. Consequently, the growth rate dependence of fermentation flux in pyruvate should present a quite different curve from that of Group A carbon sources (see Eqs. 5 and S105), which is fully validated by experimental results (Holms, 1996) (Fig. 3F).

Revised version: Finally, we consider the alteration of nutrient categories by switching to a non-Group A carbon source: pyruvate, which enters the metabolic network from the endpoint of glycolysis (Neidhardt et al., 1990; Nelson et al., 2008). The coarse-grained model for pyruvate utilization is shown in Fig. 3E (see also Fig. 1A), which shares identical precursor pools with those for Group A carbon sources, yet has several differences in the coarse-grained reactions. The growth rate dependencies of both the proteome efficiencies (see Appendix-fig. 2H) and energy fluxes (see Fig. 3F) are qualitatively similar to those of Group A carbon source utilization, while there are quantitative differences in the coarse-grained parameters (see Appendices 4.1 and 8 for details). Most notably, the critical growth rate λC(py) and the ATP production per glucose in the fermentation pathway βf(py) for pyruvate utilization are noticeably smaller than those for Group A sources (i.e., λC and βf(A) , respectively). Consequently, the growth rate dependence of fermentation flux in pyruvate should present a distinctly different curve from that of Group A carbon sources (see Eqs. 5 and S105), which is fully validated by experimental results (Holms, 1996) (see Fig. 3F).

Enzyme allocation under perturbations

Earlier version: As mentioned above, our coarse-grained model is topologically identical to the central metabolic network (Fig. 1A), and thus it can predict enzyme allocation for each gene in glycolysis and the TCA cycle (see Appendix-fig. 1B and Appendix-table 1) under various types of perturbations. In Fig. 1B, the intermediate nodes M1, M2, M3, M4, and M5 represent G6P, PEP, acetyl-CoA, α-ketoglutarate, and oxaloacetate, respectively. Then, ϕ1 and ϕ2 correspond to enzymes of glycolysis (or at the junction of glycolysis and the TCA cycle), while ϕ3 and ϕ4 correspond to enzymes in the TCA cycle.

Revised version: As mentioned above, our coarse-grained model is topologically identical to the central metabolic network (Fig. 1A) and can thus predict enzyme allocation for each gene in glycolysis and the TCA cycle (see Appendix-fig. 1B and Appendix-table 1) under various types of perturbations. In Fig. 1B, the intermediate nodes M1, M2, M3, M4, and M5 represent G6P, PEP, acetyl-CoA, α-ketoglutarate, and oxaloacetate, respectively. Therefore, ϕ1 and ϕ2 correspond to enzymes involved in glycolysis (or at the junction of glycolysis and the TCA cycle), while ϕ3 and ϕ4 correspond to enzymes in the TCA cycle (see Fig. 1A-B and Appendix 2.1).

Earlier version: We first consider enzyme allocation under carbon limitation by varying the nutrient type and concentration of a Group A carbon source (i.e., κA perturbation). In fact, this has been extensively studied in more simplified models (Hui et al., 2015; You et al., 2013), where the growth rate dependence of enzyme allocation under κA perturbation was generally considered to be a C-line response (Hui et al., 2015; You et al., 2013), i.e., the genes responsible for digesting carbon compounds show a linear increase in gene expression as the growth rate decreases (Hui et al., 2015; You et al., 2013). However, when it comes to enzymes catalyzing reactions between intermediate nodes, we collected experimental data from existing studies (Hui et al., 2015) and found that the enzymes in glycolysis exhibit a completely different response pattern compared to those in the TCA cycle (Appendix-fig. 3A-B). This discrepancy cannot be explained by the C-line response. To address this issue, we apply the coarse-grained model described above (Fig. 1B) to calculate the growth rate dependence of enzyme allocation for each ϕi (i = 1, 2, 3, 4) using the model settings for wild-type strains, where no fitting parameters are involved in determining the shape (see Eqs. S118–S119 and Appendix 8). In Figs. 4A-B and Appendix-fig. 3C-D, we see that the model predictions overall match with the experimental data (Hui et al., 2015) for representative genes from either glycolysis or the TCA cycle, and maintenance energy (with w0 = 2.5 (h-1)) has a negligible effect on this process. Still, there are minor discrepancies that arise from the basal expression of metabolic genes, which may be attributed to the fact that our model deals with relatively stable growth conditions while microbes need to be prepared for fluctuating environments (Basan et al., 2020; Kussell and Leibler, 2005; Mori et al., 2017).

Revised version: We first consider enzyme allocation under carbon limitation by varying the nutrient type and concentration of a Group A carbon source (i.e., κA perturbation). This has been extensively studied in more simplified models (Hui et al., 2015; You et al., 2013), where the growth rate dependence of enzyme allocation under κA perturbation is generally described by a C-line response (Hui et al., 2015; You et al., 2013). Specifically, the genes responsible for digesting carbon compounds exhibit a linear increase in gene expression as the growth rate decreases (Hui et al., 2015; You et al., 2013). However, when it comes to enzymes catalyzing reactions between intermediate nodes, we gathered experimental data from existing studies (Hui et al., 2015) and found that the enzymes in glycolysis exhibit a completely different response pattern compared to those in the TCA cycle (Appendix-fig. 3A-B). This discrepancy cannot be explained by the C-line response. To address this issue, we apply the coarse-grained model described above (see Fig. 1B) to calculate the growth rate dependence of enzyme allocation for each ϕi(i = 1, 2, 3, 4) using model settings for wild-type strains, with no fitting parameters influencing the shape (see Eqs. S118–S119 and Appendix 8). In Fig. 4A-B and Appendix-fig. 3C-D, we see that the model predictions overall match with the experimental data (Hui et al., 2015) for representative genes from either glycolysis or the TCA cycle, and maintenance energy (with w0 = 2.5 (h-1)) has a negligible effect on this process. Still, there are minor discrepancies that arise from the basal expression of metabolic genes, which may be attributed to the fact that our model deals with relatively stable growth conditions while microbes need to be prepared for fluctuating environments (Basan et al., 2020; Kussell and Leibler, 2005; Mori et al., 2017).

Earlier version: Relative protein expression of central metabolic enzymes under κA and ϕZ perturbations.

Revised version: Relative protein expression of central metabolic enzymes in E. coli under nutrient limitation and proteomic perturbation.

Earlier version: (A, C) Relative protein expression of representative genes from glycolysis. (B, D) Relative protein expression of representative genes from the TCA cycle. (A, B) Results of κA perturbation (Eq. S119). (C, D) Results of ϕZ perturbation (Eq. S121).

Revised version: (A, C) Relative protein expression of representative genes from glycolysis. (B, D) Relative protein expression of representative genes from the TCA cycle. (A, B) Results of the perturbation through changes in nutrient quality κA (see Eq. S119). (C, D) Results of proteomic perturbation via varied levels of expression of the useless protein LacZ (i.e., ϕZ perturbation; see Eq. S121).

Earlier version: We proceed to analyze the influence of ϕZ perturbation and energy dissipation. In both cases, our model predicts a linear response to the growth rate reduction for all genes in either glycolysis or the TCA cycle (see Appendix 5.2–5.3 for details). For ϕZ perturbation, all predicted slopes are positive, and there are no fitting parameters involved (Eqs. S120–S121). In Figs. 4C-D and Appendix-fig. 3E-J, we show that our model quantitatively illustrates the experimental data (Basan et al., 2015) for representative genes in the central metabolic network, and there is a better agreement with experiments (Basan et al., 2015) by incorporating the maintenance energy (with w0 = 2.5 (h-1)). For energy dissipation, however, the predicted slopes of the enzymes corresponding to ϕ4 are surely negative, and there is a constraint that the slope signs of the enzymes corresponding to the same ϕi (i = 1, 2, 3) should be the same. In Appendix-fig. 3K-N, we see that the model results (Eqs. S127 and S123) are consistent with experiments (Basan et al., 2015).

Revised version: We proceed to analyze the influence of ϕZ perturbation and energy dissipation. In both cases, our model predicts a linear response to growth rate reduction for all genes in either glycolysis or the TCA cycle (see Appendix 5.2–5.3 for details). For ϕZ perturbation, all predicted slopes are positive, and there are no fitting parameters involved (Eqs. S120–S121). In Fig. 4C-D and Appendix-fig. 3E-J, we show that our model quantitatively illustrates the experimental data (Basan et al., 2015) for representative genes in the central metabolic network, and there is better agreement with experiments (Basan et al., 2015) by incorporating the maintenance energy (with w0 = 2.5 (h-1) as aforementioned). For energy dissipation, however, the predicted slopes of the enzymes corresponding to ϕ4 are negative, and there is a constraint that the slope signs of the enzymes corresponding to the same ϕi(i = 1, 2, 3) should be the same. In Appendix-fig. 3K-N, we see that the model results (Eqs. S127 and S123) are consistent with experiments (Basan et al., 2015).

Added in the revised version: Explanation of the Crabtree effect in yeast and the Warburg effect in cancer cells

Added in the revised version: We proceed to apply our model to explain the Crabtree effect in yeast (Bagamery et al., 2020; De Deken, 1966; Shen et al., 2024) and the Warburg effect in tumors (Bartman et al., 2023; Duraj et al., 2021; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Shen et al., 2024; Vander Heiden et al., 2009) with slight modifications using the optimal growth principle combined with cell heterogeneity (see Appendix 9 and Appendix-fig. 5). For yeast and tumors, similar to the case of E. coli, the proteome efficiencies εr and εf are both increasing functions of the nutrient quality κA (see Eq. S170). Specifically, under poor nutrient conditions (i.e., εA is small), the proteome efficiency in respiration is higher than that in fermentation: εr > εf (see Eqs. S174–S175), making respiration the best choice for optimal growth (see Eq. S171). Conversely, when the nutrient is abundant and εf > εr, aerobic glycolysis (i.e., fermentation) becomes the optimal growth strategy (see Eq. S172). Further combination with cell heterogeneity results in the standard picture of overflow metabolism, which has indeed been observed in yeast (van Hoek et al., 1998). However, it remains challenging to tune the growth rate of cancer cells in vivo.

Added in the revised version: Recently, Shen et al. (Shen et al., 2024) discovered that the proteome efficiency measured at the cell population level in respiration (i.e., 〈εr〉; where “〈 〉” denotes the population average) is higher than that in fermentation (i.e., 〈εf〉) for many yeast and cancer cells, despite the presence of fermentation fluxes through aerobic glycolysis. Evidently, this finding (Shen et al., 2024) contradicts prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019), which hold that overflow metabolism arises because the proteome efficiency in fermentation is consistently higher than in respiration. Nevertheless, our model may resolve this puzzle due to the incorporation of two important features. First, our model predicts that the proteome efficiency in respiration is larger than that in fermentation when nutrient quality is low (see Eqs. S174–S175). Second, and crucially, by accounting for cell heterogeneity, our model allows a proportion of cells to have a higher proteome efficiency in fermentation than in respiration, even when the overall proteome efficiency in respiration at the cell population level is greater than that in fermentation (i.e., 〈εr〉 > 〈εf〉).

Added in the revised version: To compare our model results quantitatively with experimental data on yeast and tumors (Shen et al., 2024), we define Prf≡Jf(E)Jf(E)+Jr(E) as the fraction of ATP produced through fermentation. To account for cell heterogeneity, we apply Gaussian distributions to enzyme turnover numbers, as mentioned above. This yields the relationship between Prf(i.e., Jf(E)Jf(E)+Jr(E)) and 〈εr〉 and 〈εf〉 through derivations (see Eqs. S180–S190 and Appendix 9 for details): Jf(E)Jf(E)+Jr(E)=12[erf(1−〈εr〉/〈εf〉2⋅χεr2+χεf2⋅(〈εr〉/〈εf〉)2)+1], (6)

Added in the revised version: where χεr and χεf represent the CVs of proteome efficiencies εr and εf, respectively. Due to the higher levels of cell heterogeneity in yeast (Bagamery et al., 2020) and cancer cells (Duraj et al., 2021; Hanahan and Weinberg, 2011; Hensley et al., 2016), the CVs of εr and εf (i.e., χεr and χεf) in these cells are expected to be significantly higher than those in E. coli, although their precise values are unknown. The values for the variables shown in Eq. 6 can be obtained from experiments. Therefore, we plot the theoretical results from Eq. 6 using χεr and χεf values of 0.25, 0.40, and 0.58 to compare with experimental data from yeast and in vivo mouse tumors (Bartman et al., 2023; Shen et al., 2024). As shown in Fig. 5A-B, the theoretical results with χεr = χεf = 0.58 align quantitatively with the experimental data (Bartman et al., 2023; Shen et al., 2024) on both logarithmic and linear scales, demonstrating that our model has the potential to quantitatively explain the Crabtree effect in yeast and the Warburg effect in cancer cells.

Added in the revised version: Model comparison with data on the Crabtree effect in yeast and the Warburg effect in tumors.

Added in the revised version: (A) A linear scale representation on the y-axis. (B) A log scale representation on the y-axis. In (A-B), 〈εr〉 and 〈εf〉 represent the population averages of εr and εf, while χεr and χεf are the coefficients of variation (CVs) of εr and εf · 〈εr〉/〈εf〉 represents the ratio of proteome efficiency between respiration and fermentation at the population-averaged level, while Jf(E)/(Jf(E)+Jr(E)) stands for the fraction of energy flux generated by the fermentation pathway (see Eq. 6). The data for yeast in batch culture and chemostat were calculated from experimental data of S. cerevisiae and I. orientalis (Shen et al., 2024). The data for mouse tumors were calculated from in vivo experimental data of pancreatic ductal adenocarcinoma (PDAC) and leukemic spleen of mice (Bartman et al., 2023; Shen et al., 2024). See Appendix 10 for detailed information on the experimental data sources (Bartman et al., 2023; Shen et al., 2024).

Discussion

Earlier version: The phenomenon of overflow metabolism, or the Warburg effect, has been a long-standing puzzle in cell metabolism. Although many rationales have been proposed (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Molenaar et al., 2009; Niebel et al., 2019; Pfeiffer et al., 2001; Shlomi et al., 2011; Vander Heiden et al., 2009; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016) over the past century, the origin and function of this phenomenon remain unclear (DeBerardinis and Chandel, 2020; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). In this study, we use Escherichia coli as a typical example and demonstrate that overflow metabolism can be understood through optimal protein allocation combined with cell heterogeneity. In nutrient-poor conditions, the proteome energy efficiency of respiration is higher than that of fermentation (Fig. 1E), and thus the cell uses respiration to optimize growth. In rich media, however, the proteome energy efficiency of fermentation increases faster and is higher than that of respiration (Fig. 1E), leading the cell to use fermentation to accelerate growth. In further combination with cell heterogeneity in enzyme catalytic rates (Davidi et al., 2016; García - Contreras et al., 2012), our model quantitatively illustrates the thresholdanalog response (Basan et al., 2015; Holms, 1996) in overflow metabolism (Fig. 1C).

Revised version: The phenomenon of overflow metabolism, or the Warburg effect, has been a long-standing puzzle in cell metabolism. Although many rationales have been proposed over the past century (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Molenaar et al., 2009; Niebel et al., 2019; Peebo et al., 2015; Pfeiffer et al., 2001; Shlomi et al., 2011; Vander Heiden et al., 2009; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016; Zhuang et al., 2011), contradictions persist (Shen et al., 2024), leaving the origin and function of this phenomenon unclear (DeBerardinis and Chandel, 2020; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). In this study, we use E. coli as a typical example and demonstrate that overflow metabolism can be understood through optimal protein allocation combined with cell heterogeneity. Under nutrient-poor conditions, the proteome efficiency of respiration is higher than that of fermentation (see Fig. 1E), and thus the cell uses respiration to optimize growth. In rich media, however, the proteome efficiency of fermentation increases more rapidly and surpasses that of respiration (see Fig. 1E), leading the cell to adopt fermentation as the optimal growth strategy. In further combination with cell heterogeneity in enzyme catalytic rates (Davidi et al., 2016; Garcia-Contreras et al., 2012), our model quantitatively illustrates the threshold-analog response (Basan et al., 2015; Holms, 1996) in overflow metabolism (see Fig. 1C). Furthermore, it quantitatively explains the data on the Crabtree effect in yeast and the Warburg effect in cancer cells (Bartman et al., 2023; Shen et al., 2024).

Added in the revised version: In existing rationales (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Shlomi et al., 2011; Varma and Palsson, 1994; Vazquez and Oltvai, 2016), the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998) has primarily been illustrated by a threshold-linear response, which heavily depends on the assumption that cells optimize their growth rate based on a given rate of carbon influx under each nutrient condition (or similar equivalents, see Appendix 6.3). However, in practice, for microbes or tumor cells grown in vitro or in vivo, the given factors are the identity and concentration of the nutrient (Molenaar et al., 2009; Scott et al., 2010; Wang et al., 2019), rather than the rate of carbon influx. Additionally, prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019) suggest that overflow metabolism originates from the proteome efficiency in fermentation always being higher than that in respiration (see Appendix 6.3 for details). While it has been observed in E. coli that proteome efficiency in fermentation is higher than that in respiration for cells cultured in lactose at saturated concentration (Basan et al., 2015), Shen et al. (Shen et al., 2024) reported that for many yeast and cancer cells, the proteome efficiency in fermentation is noticeably lower than that in respiration, despite the presence of aerobic glycolytic fermentation flux. This observation (Shen et al., 2024) evidently contradicts the prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019). Our model resolves this puzzle by significantly differing from existing rationales in its optimization principle, where we purely optimize cell growth rate without imposing a special constraint on carbon influx (see Appendix 6.3 for details). More importantly, our model incorporates cell heterogeneity, which is crucial for both explaining the threshold-analog response in overflow metabolism and for resolving this puzzle raised by Shen et al. (Shen et al., 2024).

Earlier version: Cell heterogeneity is crucial for the threshold-analog response in overflow metabolism. In the homogeneous case, the optimal solution is a digital response (Eq. S44) that corresponds to an elementary flux mode (Müller et al., 2014; Wortel et al., 2014) and agrees with the numerical study of Molenaar et al. (Molenaar et al., 2009). However, this digital response is incompatible with the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006). By incorporating heterogeneity in enzyme catalytic rates (Davidi et al., 2016; García - Contreras et al., 2012), the critical growth rate (i.e., threshold) changes from a single value into a Gaussian distribution (Eq. 45, see Appendix 7 for details; see also Appendix-fig. 4) for a cell population, thus turning a digital response into the threshold-analog response in overflow metabolism (Fig. 1C). Our model results relying on cell heterogeneity are fully validated by the observed distributions of single-cell growth rate (Wallden et al., 2016) (Appendix-fig. 2B) and experiments with various types of perturbations (Basan et al., 2015; Holms, 1996; Hui et al., 2015), both for acetate secretion patterns and gene expression in the central metabolic network (Figs. 2–4 and Appendix-figs. 2D-E and 3).

Revised version: In the homogeneous case, the optimal growth strategy for growth rate dependent fermentation flux results in a digital response (see Eq. S44), corresponding to an elementary flux mode (Müller et al., 2014; Wortel et al., 2014), which aligns with the numerical study by Molenaar et al. (Molenaar et al., 2009), yet is incompatible with the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998). Furthermore, cells would not generate fermentation flux if the proteome efficiency in fermentation were higher than that in respiration under the optimal growth framework. By incorporating heterogeneity in enzyme catalytic rates (Davidi et al., 2016; Garcia-Contreras et al., 2012), the critical growth rate (i.e., threshold) shifts from a single value to a Gaussian distribution (see Eq. S45 and Appendix 7 for details; see also Appendix-fig. 4) across a cell population, thereby turning a digital response into the threshold-analog response observed in overflow metabolism (see Fig. 1C). Moreover, cell heterogeneity allows a fraction of cells to possess a larger proteome efficiency in fermentation than in respiration despite the overall proteome efficiency in respiration at the cell population level is higher than in fermentation. This mechanism facilitates the fermentation flux in yeast and cancer cells observed by Shen et al. (Shen et al., 2024) (see Fig. 5A-B).

Added in the revised version: Our model results, based on cell heterogeneity, are further supported by observed distributions of single-cell growth rates in E. coli (Wallden et al., 2016) (see Appendix-fig. 2B), and by experiments involving various types of perturbations (Basan et al., 2015; Holms, 1996; Hui et al., 2015), both in terms of acetate secretion patterns and gene expression in the central metabolic network (see Fig. 2–4, Appendix-figs. 2D-E and 3). Furthermore, the heterogeneity patterns predicted by our model for fermentation and respiration modes within a cell population are validated by experiments on intra-tumor heterogeneity in glioblastoma using optical microscopy (Duraj et al., 2021) and align with the non-genetic heterogeneity observed in experiments with S. cerevisiae (Bagamery et al., 2020) and E. coli (Nikolic et al., 2013). Finally, our model can be broadly applied to address heterogeneity-related challenges in metabolism on a quantitative basis, including diverse metabolic strategies of cells in various environments (Bagamery et al., 2020; Duraj et al., 2021; Escalante-Chong et al., 2015; Hensley et al., 2016; Liu et al., 2015; Solopova et al., 2014; Wang et al., 2019).

Removed in the revised version: Finally, our model can be broadly used to address heterogeneity-related challenges in metabolism on a quantitative basis, including the Crabtree effect in yeast (Bagamery et al., 2020; De Deken, 1966), the Warburg effect in cancer (Duraj et al., 2021; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009) (see Appendix 6.4 for an explanation of the Warburg effect), and the heterogeneous metabolic strategies of cells in various types of environments (Bagamery et al., 2020; Duraj et al., 2021; Escalante-Chong et al., 2015; Hensley et al., 2016; Liu et al., 2015; Solopova et al., 2014; Wang et al., 2019).